# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In this GitHub repository, there is a file called `allergies.json` that contains a list of patient allergies.  You will need to download this [file from here](https://raw.githubusercontent.com/paulboal/hds5210-2023/main/week06/allergies.json) and then upload it into Google Colab to run these examples. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE="/allergies.json"
""" jason file will be copied to the ALLERGIES_FILE which will be used in the code"""
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
from typing_extensions import Required
# Put your solution here
import json
def allergy_count(json_file):
    try:
        """In this try block json will be loaded as the read only file
           And the file will be loaded to the data required"""
        with open(json_file, 'r') as file:
            data_required = json.load(file)
            """We have to give the entry which is a list of things and
               it will return lengthnof the data that is required"""
            if isinstance(data_required.get("entry"), list):
                return len(data_required["entry"])
                """If the file does not contain an entry it
                  will return the it does not contain the list of entries"""
            else:
                print("Error: The JSON file does not contain a list of entries.")
                return 0
                """If we does not the load the file properly
                it will raise the error as file was not found"""
    except FileNotFoundError:
        print(f"Error: File '{json_file}' not found.")
        return 0
    except json.JSONDecodeError:
        print(f"Error: Unable to decode JSON in '{json_file}'.")
        return 0
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
# @title Default title text
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
import json
ALLERGIES_FILE="allergies.json"
```

```python
# Put your solution here
import json
ALLERGIES_FILE="/allergies.json"
import json
def patient_count(json_file):
    try:
      """In this try block json will be loaded as the read only file
           And the file will be loaded to the data required"""
      with open(json_file, 'r') as file:
            data = json.load(file)
            entries = data.get("entry", [])
            """By writing this statement entries = data.get("entry", [])
            it will return the list of unique patients in the allergies.json"""
            unique_patients = set()
            for entry in entries:
                """In this loop we will get the resource and from that we will get details of the patient
                   and from that we will dispaly the patient details"""
                patient_display = entry.get("resource", {}).get("patient", {}).get("display")
                if patient_display:
                    unique_patients.add(patient_display)
                """And finally it will return the length of unique patients list"""
            return len(unique_patients)
            """If we does not the load the file properly
                it will raise the error as file was not found"""
    except FileNotFoundError:
        print(f"Error: File '{json_file}' not found.")
        return 0
    except json.JSONDecodeError:
        print(f"Error: Unable to decode JSON in '{json_file}'.")
        return 0
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
import json
ALLERGIES_FILE="/allergies.json"
```

```python
# Put your solution here
def allergy_per_patient(ALLERGIES_FILE):
    allergies_by_patient = {}
    with open(ALLERGIES_FILE, 'r') as file:
        """In this try block json will be loaded as the read only file
           And the file will be loaded to the data required"""
        data = json.load(file)
        for entry in data.get('entry', []):
            """From this we can get resources that are related to the patiets"""
            resource = entry.get('resource', {})
            patient = resource.get('patient', None)
            if patient is not None:
                # Extract the patient name from the dictionary
                patient_name = patient.get('display', None)
                if patient_name:
                    """In this we will get deatils of the allergies that a patient had
                        an dit wil return the data"""
                    if patient_name in allergies_by_patient:
                        allergies_by_patient[patient_name] += 1
                    else:
                        allergies_by_patient[patient_name] = 1
    return allergies_by_patient
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
import json

def allergy_list(json_file):
    try:
        """In this try block json will be loaded as the read only file
           And the file will be loaded to the data required"""
        with open(json_file, 'r') as file:
            data = json.load(file)
            entries = data.get("entry", [])
            """In this we will get the entries of the patients
               for which we have to create the data entry"""
            output_list = []
            for entry in entries:
                """ In thi loop it will return thempatient name, substance text,representation
                    and a reaction manifestation for the each and every patient"""
                patient_name = entry.get("resource", {}).get("patient", {}).get("display")
                substance_text = entry.get("resource", {}).get("substance", {}).get("text")
                reaction_text = entry.get("resource", {}).get("reaction", [{}])[0].get("manifestation", [{}])[0].get("text")
                """In this it will append all the details to the output list of th epatients"""
                if patient_name and substance_text and reaction_text:
                    output_list.append([patient_name, substance_text, reaction_text])

            return output_list
            """If we does not the load the file properly
                it will raise the error as file was not found"""
    except FileNotFoundError:
        print(f"Error: File '{json_file}' not found.")
        return []
    except json.JSONDecodeError:
        print(f"Error: Unable to decode JSON in '{json_file}'.")
        return []

```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_reaction(json_file, patient, substance):
    allergy_entries = allergy_list(json_file)
    """ In this allergy_reaction function we will get the reaction when we
        eneter the details of the patient with their name and the substance"""
    for entry in allergy_entries:
        entry_patient, entry_substance, reaction = entry
        if entry_patient == patient and entry_substance == substance:
            return reaction

    return None

```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

```python
from google.colab import drive
drive.mount('/content/drive')
```

Output:
```
Mounted at /content/drive
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

